// variables1.rs
//
// Make me compile!
//
// Execute `rustlings hint variables1` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

// fn main() {
//     let x:u32 = 5;
//     println!("x has the value {x}!");
// }

// variables2.rs
//
// Execute `rustlings hint variables2` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

// fn main() {
//     let x;
//     if x == 10 {
//         println!("x is ten!");
//     } else {
//         println!("x is not ten!");
//     }
// }

// variables3.rs
//
// Execute `rustlings hint variables3` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

// fn main() {
//     let x: i32 = 50;
//     println!("Number {}", x);
// }

// variables4.rs
//
// Execute `rustlings hint variables4` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

// fn main() {
//     let mut x = 3;
//     println!("Number {}", x);
//     x = 5; // don't change this line
//     println!("Number {}", x);
// }


